```python
# Initialize Otter
import otter
grader = otter.Notebook("p1.ipynb")
```

```python
# PLEASE FILL IN THE DETAILS
# Enter None if you don't have a project partner

# project: p1
# submitter: ratushko
# partner: ppopuri
# hours: 1
```

**Question 1:** "Hello, World!" program.

```python
name = "World"
greeting = "Hello, " + name + "!"
greeting
```

Output:
```
'Hello, World!'
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Compute current year.

```python
curr_year = 45 * 45 - 3
curr_year
```

Output:
```
2022
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

SUBMISSION INSTRUCTIONS: The zipfile automatically gets downloaded. You should have received GradeScope invite - please accept it. Login to GradeScope and upload the zipfile. Check otter results as soon as auto-grader gets completed.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...

Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!

    q2 results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q1 (points: 12.5)
```
>>> greeting == "Hello, World!"
True
```


Tests for q2 (points: 12.5)
```
>>> curr_year == 2022
True
```
